```python
# Initialize Otter
import otter
grader = otter.Notebook("intro2Seaborn.ipynb")
```

## Lecture Section

In this lecture, we will cover the `seaborn` library. `seaborn` is a statistical plotting library that is built on-top of matplotlib. It creates beautiful plots with minimal effort.

We will cover:
* Customizing Plots
    * Scaling
    * Themes
    * Color Palettes
* Seaborn Specific Plots
    * Violinplot
    * Swarmplot
    * Histplot
    * Distplot
    * Jointplot
    * Pairplot

This lecture-assignment is completely graded by-hand, there is no autograder. To maintain consistency, it used the same technology as the autograded lecture-assignments. comman

We will use the `iris` dataset because it is clean and easily accessible.

```python
import seaborn as sns
import matplotlib.pyplot as plt

data = sns.load_dataset("iris")
data

```

### Customization

There are five themes we can use when plotting with seaborn. They are `white`, `dark`, `whitegrid`, `darkgrid`, and `ticks`. Try a couple below.

```python
sns.lineplot(x="sepal_length", y="sepal_width", data=data)
sns.set_style("whitegrid")
plt.show()
```

We can easily scale our figure using `set_context()`- and options like `paper`, `talk`, and `poster`.

```python
sns.lineplot(x="sepal_length", y="sepal_width", data=data)
sns.set_context("poster")

plt.show()
```

Color palettes are one of the most important features of seaborn.

To check the current color palette we are using, we call `sns.palette()`

```python
sns.color_palette()

```

The color palette above is a **qualitative** color palette. It works well with categorical data.

The palette below is a **diverging** color palette. It has colors that can be easily told apart on each end, and a nuetral color in the middle.

```python
sns.color_palette('PiYG', 11)

```

The last type of color palette is **sequential**. This type of palette focuses on one color and brightens from left to right.

```python
sns.color_palette('Greens', 11)

```

There are a lot of built-in palettes. View them here:
https://seaborn.pydata.org/tutorial/color_palettes.html#sequential-color-palettes
or
https://www.practicalpythonfordatascience.com/ap_seaborn_palette

For all palettes, the `_r` version is a continuous color gradiant, while the name without `_r` is a discrete version.

```python
sns.color_palette("rocket_r", as_cmap=True)
```

If we wish to change the default color palette, we use `set_palette()`

```python
data = sns.load_dataset("iris")
sns.set_palette('rocket')
sns.lineplot(x="sepal_length", y="sepal_width", data=data)

```

### Seaborn Specific Plots

We will briefly talk about a few unique plots. These plots may not be exclusive to `seaborn` but they do look nicer. To save time, we won't cover the simpler plots that we did with `matplotlib` because they work nearly the same. Don't forget that they exist, though!

**Violinplot** : similar to a boxplot, but uses the kernel density estimation to graph the shape. Thick sections show a higher distribution of points, and thinner sections show a lower distribution of points. THe white, centered dot is the median, and the thicker line protruding from it is the interquartile range. Finally, the thin line in the middle is 1.5x interquartile range. Points outside of this range are outliers.

```python
data = sns.load_dataset("iris")
sns.set_palette('bright')
sns.violinplot(x='species', y='sepal_width', data=data, hue = 'species') # hue colors based on category
```

**Swarmplot** : similar to violinplot, but with less information and dots. It can be considered a scatterplot for categories.

```python
data = sns.load_dataset("iris")
sns.set_palette('bright')
sns.swarmplot(x='species', y='sepal_width', data=data, hue = 'species') # hue colors based on category
```

Some people like them together...

```python
data = sns.load_dataset("iris")
sns.set_palette('dark')
sns.violinplot(x='species', y='sepal_width', data=data, hue = 'species') # hue colors based on category
sns.set_palette('hls', 3)
sns.swarmplot(x='species', y='sepal_width', data=data, hue = 'species') # hue colors based on category
```

**Histplot** : a histplot is a histogram-like plot for categorical data. Each column shows the distribution of data for that specific category.

```python
data = sns.load_dataset("iris")
sns.histplot(x='species', y='sepal_width', data=data)
```

We can give `histplot()` one column for a normal histogram with counts, and we can graph the `kde` if we choose.

```python
sns.histplot(data["sepal_width"], kde=True)

```

One step further, we can turn the counts into a density.

```python
sns.histplot(data["sepal_width"], kde=True, stat='density')
```

**Jointplot** : a jointplot plots two variables with univariate and bivariate graphs

```python
data = sns.load_dataset("iris")
sns.jointplot(x='species', y='sepal_width', data=data)
```

**Pairplot** : a pairplot graphs a pairwise relationship. It creates a plot between each numerical column. The diagonal has the density plot for each variable. All other plots are scatterplots between the two variables. Note that this is a little different that the usual `ggpairs()` method in R: we do not have a correlation coefficient in the upper triangle here.

```python
data = sns.load_dataset("iris")
sns.pairplot(data=data, hue='species')
```

Pairplots can be very computationally expensive. I recommend saving the figure as an image, then displaying that image instead of re-running the code (unless your code has changed).

## Assignment Section

Some of the questions will be manually graded, and some will be autograded.


<!-- BEGIN QUESTION -->

**Question 1.**

Use the 'tips' dataset from seaborn. It is loaded in for you.

Create a histogram of the total_bill variable.
Make sure the plot has an appropriate title and axis labels. Do not forget to show the plot!

```python
import seaborn as sns
import matplotlib.pyplot as plt

tips = sns.load_dataset("tips")
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.**

Create a scatterplot of total_bill vs. tip.
Then add a linear regression line to the plot. Color points by sex.

```python
import seaborn as sns
import matplotlib.pyplot as plt

tips = sns.load_dataset("tips")
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3.**

```python
import seaborn as sns
import matplotlib.pyplot as plt

tips = sns.load_dataset("tips")
...
```

<!-- END QUESTION -->

**Question 4.**

Please ignore me :) - I am a free point!

```python
x = 1
```

```python
grader.check("q4")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q4 (points: None)
```
>>> assert x == 1
```
